# Week 2 Problem Set

## Cohort Sessions


```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
Version 1.0.5
```

**CS1.** *Binary Heap:* Write the following functions:



- `parent_of(index)`: returns the index of node's parent
- `left_of(index)`: returns the index of node's left child
- `right_of(index)`: returns the index of node's right child
- `max_child(array, index, heap_size)`: returns the index of node's largest child. You can assume that the node has at least one child.

Hint: 
- `index` starts from 0. 
- You can refer to the pseudocode in [Binary Heap and Heapsort](https://data-driven-world.github.io/2023/notes/category/week-2-analysing-programs).
- When finding the index of the largest child, consider the following cases:
    1. when the node only has one child
    2. when the node has two children


```python
def parent_of(index: int) -> int:
    return int((index-1)/2)
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert parent_of(1) == 0 
assert parent_of(2) == 0
assert parent_of(5) == 2
assert parent_of(6) == 2

```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def left_of(index: int) -> int:
    return int((index*2)+1)
```

<!-- grade_id: cs13, points: 1, grade: true, solution: false -->
```python
assert left_of(0) == 1
assert left_of(1) == 3
assert left_of(3) == 7
assert left_of(6) == 13

```

<!-- grade_id: cs14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def right_of(index: int) -> int:
    return int((index+1)*2)
```

<!-- grade_id: cs15, points: 1, grade: true, solution: false -->
```python
assert right_of(0) == 2
assert right_of(1) == 4
assert right_of(3) == 8
assert right_of(5) == 12

```

<!-- grade_id: cs16, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Re-order the code for returning the index of the max child from a given node index and some particular heap size.

Input:
- array: containing the binary heap data
- index: node index which we want to find the max child
- heap_size: the number of elements in the heap

Note that the array size can be bigger than the heap size. In this case, ignore the rest of the element as NOT part of the heap.

Output:
- index of the max child

Write the code below and test.

```python
def max_child(array: list[int|float], index: int, heap_size: int) -> int:
    # If the right child index is out of the heap size, return the left child index
    if right_of(index) >= heap_size:
        return left_of(index)
    else:
        # Compare the values of the left and right children and return the index of the larger one
        if array[left_of(index)] > array[right_of(index)]:
            return left_of(index)
        else:
            return right_of(index)
```

<!-- grade_id: cs17, points: 1, grade: true, solution: false -->
```python
###  AUTO-GRADED TEST  ###
```

<!-- grade_id: cs18, points: 1, grade: true, solution: false -->
```python
###  AUTO-GRADED TEST  ###
```

**CS2.** *Binary Heap:* Write two functions 

- `max_heapify(array, index, size)`: that moves the node down so as to satisfy the heap property.  The first argument is the array that contains the heap. The second argument is an integer index where to start the process of heapifying. The third argument is the size of the heap in the array. This argument will be useful in heapsort algorithm where we take out the elements in the array from the heap. 

Hint: You should make use of `size` argument to determine the last element of the heap in the array rather than `len(array)`.

- `build_max_heap(heap)`: that builds the max heap from any array. This function should make use of `max_heapify()` in its definition.

Hint: You can refer to the pseudocode in [Binary Heap and Heapsort](https://data-driven-world.github.io/2023/notes/category/week-2-analysing-programs) for the above functions.

```python
def max_heapify(array: list[int|float], index: int, size: int) -> None:
    current_i = index
    # Continue until the current node has a left child within the heap size
    while left_of(current_i) < size:
        # Find the index of the largest child
        max_child_i = max_child(array, current_i, size)
        # If the largest child is greater than the current node, swap them
        if array[max_child_i] > array[current_i]:
            array[max_child_i], array[current_i] = array[current_i], array[max_child_i]
            # Move to the largest child index
            current_i = max_child_i
        else:
            # If the current node is larger than both children, the heap property is satisfied
            break
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
###  AUTO-GRADED TEST  ###
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###  AUTO-GRADED TEST  ###
```

```python
def build_max_heap(array: list[int|float]) -> None:
    # Get the length of the array
    n = len(array)
    # Calculate the starting index for heapifying, which is the parent of the last element
    starting_index = n // 2 - 1
    # Iterate from the starting index to the root (index 0)
    for current_index in range(starting_index, -1, -1):
        # Heapify the subtree rooted at the current index
        max_heapify(array, current_index, n)
```

<!-- grade_id: cs23, points: 1, grade: true, solution: false -->
```python
###  AUTO-GRADED TEST  ###
```

<!-- grade_id: cs24, points: 1, grade: true, solution: false -->
```python
###  AUTO-GRADED TEST  ###
```

**CS3.** *Heapsort:* Implement heapsort algorithm following the pseudocode in [Binary Heap and Heapsort](https://data-driven-world.github.io/2023/notes/category/week-2-analysing-programs).



```python
def heapsort(array: list[int|float]) -> None:
    # Build a max heap from the input array
    build_max_heap(array)
    # Initialize the position of the end of the heap
    heap_end_pos = len(array) - 1
    # Continue until the heap is reduced to one element
    while heap_end_pos > 0:
        # Swap the root (maximum value) with the last element of the heap
        array[0], array[heap_end_pos] = array[heap_end_pos], array[0]
        # Reduce the size of the heap by one
        heap_end_pos -= 1
        # Restore the heap property for the reduced heap
        max_heapify(array, 0, heap_end_pos + 1)
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
###  AUTO-GRADED TEST  ###
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###  AUTO-GRADED TEST  ###
```

**CS4.** Measure computational time of Python's built-in sort function by filling the template below. Hint: 
- You will need the function `gen_random_int()` from Week 01 Problem Set.
- Use `sorted(list)` function of Python's list See [Python's Sorting HOW TO Documentation](https://docs.python.org/3.3/howto/sorting.html). 

```python
import time
import random
from collections.abc import Callable

def run_function(f: Callable[[list[int]], list[int]], x: list[int]) -> float:
    start: float = time.time()
    f(x)
    end: float = time.time()
    return end-start

def gen_random_int(number: int, seed: int) -> list[int]:
    result = []
    for i in range(number):
        result.append(i)
        
    random.Random(seed).shuffle(result)
    return result

 
time_builtin: list[float] = []
# set the maximum power for 10^power number of inputs
maxpower: int = 7

for n in range(1, maxpower + 1):
    # create array for 10^1, 10^2, etc 
    # use seed 100
    array = gen_random_int(10**n,100)
    # call run_function with sorted() and array as arguments
    result = run_function(sorted, array)
    time_builtin.append(result)


print(time_builtin)
```

Output:
```
[0.0, 0.0, 0.0, 0.0010004043579101562, 0.015000343322753906, 0.37488317489624023, 4.940637826919556]
```