# Week 4 Problem Set

## Cohort Sessions

```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
Version 1.0.5
```

```python
from typing import TypeAlias
from typing import Optional, Any

Number: TypeAlias = int | float
```

**CS1**. Create a class called `Square`. Its initialization takes in a variable which is the side of of a square. The object has one attribute which is called `side` to represent the side of a square. It has one property called `area` which sets the dimension of the square. Setting the area changes the side of a square.

```python
class Square:
    # Initialize the square with a given side length (default is 1)
    def __init__(self, dim: Number=1) -> None:
        self.side = dim

    # Property to get the area of the square
    @property
    def area(self) -> Number:
        return self.side ** 2
   
    # Setter to set the area of the square, which adjusts the side length accordingly
    @area.setter
    def area(self, value: Number) -> None:
        self.side = value ** 0.5

    # String representation of the square
    def __str__(self) -> str:
        return f'Square with side {self.side}'

```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
s: Square = Square()
assert s.area == 1

s: Square = Square(2)
assert s.area == 4
s.area = 9
assert s.side == 3

```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** We are going to create a simple Car Racing game. First, let's create a class Car with the following properties:
- `racer` which stores the name of the driver. This property must be non-empty string. This property should be initialized upon object instantiation.
- `speed` which stores the speed of the car. This property can only be non-negative values and must be **less than or equal** to a maximum speed. The setter of this property should ensure that the value of `speed` should only falls between `0` and `max_speed` (inclusive). 
- `pos` which is an integer specifying the position of the car which can only be non-negative values.
- `is_finished` which is a computed property that returns `True` or `False` depending whether the position has reached the finish line.

Each car also has the following attributes:
- `max_speed` which specifies the maximum speed the car can have. This attribute should be initialized upon object instantiation.
- `finish` which stores the finish distance the car has to go through. Upon initialization, it should be set to -1. 

The class has the following methods:
- `start(init_speed, finish_distance)` which set the speed property to some initial value. The method also set the finish distance to some value and set the `pos` property to 0.
- `race(acceleration)` which takes in an integer value for its acceleratin and modify both the speed and the position of the car.



```python
class RacingCar:
    
    def __init__(self, name: str, max_speed: int) -> None:
        # Initialize the car with racer's name and maximum speed
        self._racer = name
        self._speed = 0
        self._pos = 0
        self.max_speed = max_speed
        self.finish = -1
        self.is_finished
    
    @property
    def racer(self) -> str:
        # Get the racer's name
        return self._racer
    
    @racer.setter
    def racer(self, name: str) -> None:
        # Set the racer's name if it's a non-empty string
        if not isinstance(name, str) or not name:
            print("Racer name cannot be empty or be an integer.")
        else:
            self._racer = name
            
    @property
    def speed(self) -> int:
        # Get the current speed of the car
        return self._speed
    
    @speed.setter
    def speed(self, val: int) -> None:
        # Set the speed of the car ensuring it is within the valid range
        if val < 0 or val > self.max_speed:
            print(f"Speed must be between 0 and {self.max_speed}.")
        else:
            self._speed = val
        
    @property
    def pos(self) -> int:
        # Get the current position of the car
        return self._pos
    
    @pos.setter
    def pos(self, val: int) -> None:
        # Set the position of the car ensuring it is non-negative
        if val < 0:
            print("Position cannot be negative.")
        else:
            self._pos = val
            
    @property
    def is_finished(self) -> bool:
        # Check if the car has finished the race
        if self.finish != -1:    
            return self._pos >= self.finish
        return False
            
    def start(self, init_speed: int, finish_dist: int) -> None:
        # Start the race with initial speed and finish distance
        self.speed = init_speed
        self.finish = finish_dist
        self._pos = 0
    
    def race(self, acc: int) -> None:
        # Update the speed and position of the car based on acceleration
        new_speed = self._speed + acc
        self.speed = max(0, min(new_speed, self.max_speed))
        self._pos += self._speed
        
    def __str__(self) -> str:
        # String representation of the car's current state
        return f"Racing Car {self.racer} at position: {self._pos}, with speed: {self._speed}."
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
car: RacingCar = RacingCar("Hamilton", 200)
assert car.racer == "Hamilton"
assert car.max_speed == 200
assert car.finish == -1

car.racer = "Bottas"
assert car.racer == "Bottas"
car.racer = ""
assert car.racer == "Bottas"
car.racer = 21
assert car.racer == "Bottas"

car.speed = 10
assert car.speed == 10
car.speed = 0
assert car.speed == 0
car.speed = -10
assert car.speed == 0
car.speed = car.max_speed
assert car.speed == car.max_speed
car.speed = car.max_speed + 10
assert car.speed == car.max_speed

car.pos = 10
assert car.pos == 10
car.pos = -10
assert car.pos == 10
car.pos = 0
assert car.pos == 0

assert not car.is_finished
car.finish = 20
car.pos = 10
assert not car.is_finished
car.pos = 20
assert car.is_finished
car.pos = 22
assert car.is_finished

car.start(50, 200)
assert car.pos == 0
assert car.speed == 50
assert car.finish == 200

car.race(0)
assert car.speed == 50
assert car.pos == 50
assert not car.is_finished

car.race(10)
assert car.speed == 60
assert car.pos == 110
assert not car.is_finished

car.race(-10)
assert car.speed == 50
assert car.pos == 160
assert not car.is_finished

car.race(0)
assert car.is_finished

```

Output:
```
<cell>10: [1m[91merror:[0m Incompatible types in assignment (expression has type [0m[1m"int"[0m, variable has type [0m[1m"str"[0m)  [0m[93m[assignment][0m
Racer name cannot be empty or be an integer.
Racer name cannot be empty or be an integer.
Speed must be between 0 and 200.
Speed must be between 0 and 200.
Position cannot be negative.
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3**. Design a class named `Stock` to represent a company's stock that contains:
* a string computed property named `symbol` for the stock's symbol.
* a string computed property named `name` for the stock's name.
* a float property named `previous_closing_price` that stores the stock's price for the previous day.
* a float property named `current_price` that stores the stock price for the current time.
* an initializer that initializes the new stock object with the specified symbol, name, previous price and current price. 
* a method named `get_change_percent()` that returns the percentage changed from previous closing price to current price. 

```python
class Stock:
    # Initialize the stock with symbol, name, previous closing price, and current price
    def __init__(self, symbol: str, name: str, previous_closing_price: float, current_price: float):
        self._symbol = symbol
        self._name = name
        self.previous_closing_price = previous_closing_price
        self.current_price = current_price

    @property
    def symbol(self) -> str:
        # Get the stock's symbol
        return self._symbol

    @property
    def name(self) -> str:
        # Get the stock's name
        return self._name

    def get_change_percent(self) -> float:
        # Calculate and return the percentage change from previous closing price to current price
        return ((self.current_price - self.previous_closing_price) / self.previous_closing_price) * 100

```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
s: Stock = Stock("INTC", "Intel Corporation", 20.5, 20.35)
assert s.symbol == "INTC"
assert s.name == "Intel Corporation"
assert s.previous_closing_price == 20.5
assert s.current_price == 20.35
assert round(s.get_change_percent(),2) == -0.73

```

**CS4.** Create a class called `Fraction` to represent a simple fraction. The class has two properties:
- `num`: which represents a numerator and of the type Integer.
- `den`: which represents a denominator and of the type Integer. Denominator should not be a zero. If a zero is assigned, you need to replace it with a `1`. 
If a `float` is assigned to either `num` or `den`, the setter should convert it to an integer using `int()` function. The setter should accept only `int` or `float` type of data.

The class should have the following method:
- `__init__(num, den)`: to initialize the numerator and the denominator. You should check if the denominator is zero. If it is you should assign `1` as the denominator instead. 
- `__str__()`:  for the object instance to be convertable to String.  You need to return a string in a format of `num/den`


```python
class Fraction:
    def __init__(self, num: int, den: int) -> None:
        # Initialize the numerator and denominator, ensuring denominator is not zero
        self.num = num
        self.den = den if den != 0 else 1
    
    @property
    def num(self) -> int:
        # Get the numerator
        return self._num
    
    @num.setter
    def num(self, val: int) -> None:
        # Set the numerator, converting float to int if necessary
        if isinstance(val, (int, float)):
            self._num = int(val)
        else:
            print("Numerator must be an integer or float")
    
    @property
    def den(self) -> int:
        # Get the denominator
        return self._den
    
    @den.setter
    def den(self, val: int) -> None:
        # Set the denominator, converting float to int if necessary and ensuring it is not zero
        if isinstance(val, (int, float)):
            val = int(val)
            if val == 0:
                self._den = 1
            else:
                self._den = val
        else:
            print("Denominator must be an integer or float")
    
    def __str__(self) -> str:
        # Return the string representation of the fraction
        return f"{self.num}/{self.den}"

```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
f0: Fraction = Fraction(0, 1)
assert f0.num == 0
assert f0.den == 1
assert str(f0) == "0/1"

f1: Fraction = Fraction(1, 2)
assert f1.num == 1
assert f1.den == 2
assert str(f1) == "1/2"

f1.num = 3
f1.den = 4
assert str(f1) == "3/4"

```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5.** Implement the `Stack` abstract data type using a Class. You can use `list` Python data type as its internal data structure. Name this `list` as `items`.

The class should have the following interface:
- `__init__()` to initialize an empty list for the stack to store the items.
- `push(item)` which stores an item into the top of the stack.
- `pop()` which returns and removes the top element of the stack. The return value is optional as it may return `None` if there are no more elements in the stack.
- `peek()` which returns the top element of the stack. If the stack is empty, it returns `None`.

The class should have the following properties:
- `is_empty` is a computed property which returns either `True` or `False` depending whether the stack is empty or not.
- `size` is a computed property which returns the number of items in the stack.


```python
class Stack:
    def __init__(self) -> None:
        # Initialize an empty list to store stack items
        self.__items: list[Any] = []

    def push(self, item: Any) -> None:
        # Add an item to the top of the stack
        self.__items.append(item)

    def pop(self) -> Any:
        # Remove and return the top item of the stack, if the stack is not empty
        if not self.is_empty:
            return self.__items.pop()
        return None

    def peek(self) -> Any:
        # Return the top item of the stack without removing it, if the stack is not empty
        if not self.is_empty:
            return self.__items[-1]
        return None

    @property
    def is_empty(self) -> bool:
        # Check if the stack is empty
        return len(self.__items) == 0

    @property
    def size(self) -> int:
        # Return the number of items in the stack
        return len(self.__items)
```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
s1: Stack = Stack()
s1.push(2)
assert not s1.is_empty
assert s1.pop() == 2
assert s1.is_empty
assert s1.pop() == None
s1.push(1)
s1.push(2)
s1.push(3)
assert not s1.is_empty
assert s1._Stack__items == [1, 2, 3]
assert s1.peek() == 3
assert s1.size == 3

```

<!-- grade_id: cs52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS6.** Implement a Queue abstract data structure using two Stacks instead of Python's list. For this double-stack implementation, the Queue has a *left* Stack and a *right* Stack. The enqueue and dequeue operations work as follows:
- enqueue: This operation just pushes the new item into the *left* Stack.
- dequeue: This operation is done as follows:
    - if the *right* Stack is empty: create a new *right* Stack which is the reverse of the items in the *left* Stack. You should then empty the *left* stack.
    - if the *right* Stack is not empty: pop from the *right* Stack.
- peek: This operation is similar to Stack's peek. It returns the value at the front of the Queue without removing it. If the Queue is empth, it should return a `None` object.

Note that, you could also enqueue to the right Stack and dequeue from the left stack. You can choose to implement according to this if you wish and make sure the operations still work as defined.

```python
class Queue:
    def __init__(self) -> None:
        # Initialize two stacks to implement the queue
        self.left_stack: Stack = Stack()
        self.right_stack: Stack = Stack()
    
    def enqueue(self, item: Any) -> None:
        # Add an item to the left stack
        self.left_stack.push(item)

    def dequeue(self) -> Any:
        # Remove and return the front item of the queue
        if self.right_stack.is_empty:
            # Transfer all items from left stack to right stack if right stack is empty
            while not self.left_stack.is_empty:
                self.right_stack.push(self.left_stack.pop())
        return self.right_stack.pop()

    def peek(self) -> Any:
        # Return the front item of the queue without removing it
        if self.right_stack.is_empty:
            # Transfer all items from left stack to right stack if right stack is empty
            while not self.left_stack.is_empty:
                self.right_stack.push(self.left_stack.pop())
        return self.right_stack.peek()
    
    @property
    def is_empty(self) -> bool:
        # Check if both stacks are empty
        return self.left_stack.is_empty and self.right_stack.is_empty
        
    @property
    def size(self) -> int:
        # Return the total number of items in both stacks
        return self.left_stack.size + self.right_stack.size
```

<!-- grade_id: cs61, points: 1, grade: true, solution: false -->
```python
q1: Queue = Queue()
q1.enqueue(2)
assert not q1.is_empty 
assert q1.size == 1
assert q1.dequeue() == 2
assert q1.is_empty
q1.enqueue(1)
q1.enqueue(2)
q1.enqueue(3)
assert q1.size == 3
assert q1.peek() == 1
assert q1.dequeue() == 1
assert q1.dequeue() == 2
assert q1.dequeue() == 3
assert q1.peek() == None

```

<!-- grade_id: cs62, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```