# Week 1 Problem Set

## Cohort Problems



```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
The nb_mypy extension is already loaded. To reload it, use:
  %reload_ext nb_mypy
```

**CS0-a.** Find a maximum in a list. Write a function called `find_max` that takes in a list and return the maximum element in that list. **Do not use the built-in `max()` function**.


```python
from typing import Optional

def find_max(array: list[int]) -> Optional[int]:
    # Check if the array is empty
    if array == []:
        return None
    else:
        # Initialize the maximum value with the first element of the array
        max_value = array[0]
        
        # Iterate through the array to find the maximum value
        for i in range(len(array)):
            if max_value < array[i]:
                max_value = array[i]
                
    return max_value

```

<!-- grade_id: cs01, points: 1, grade: true, solution: false -->
```python
assert find_max([1, 2, 3, 4, 5]) == 5
assert find_max([5, 4, 3, 2, 1]) == 5
assert find_max([1, 2, 5, 4, 3]) == 5
assert find_max([0, 0, 0, 0, 0]) == 0
assert find_max([-10, -20, -30, -40, -50]) == -10
assert find_max([1]) == 1
assert find_max([]) == None

```

<!-- grade_id: cs02, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS0-b.** Write a function `bubble_up` which takes in a list of numbers and modifies this input list in such a way that the largest element ends up in the last position. It takes in a second optional argument of an ending index in which no more comparison needs to be made. By default the value of this ending index is `len(array)` which is the index after the last element. This means that by default it should compares all the way to the last element. 

Note: The function should keeps the number of element the same and shift the lower numbers to the left if needed. 

Example:

```
array = [1, 100, 2, 3, 4, 5]
bubble_up(array)
print(array)
```

The output should be the following.
```
[1, 2, 3, 4, 5, 100]
```

Another example:
```
array = [1, 100, 2, 3, 4, 5]
bubble_up(array, 5)
print(array)
```

The output should be the following.

```
[1, 2, 3, 4, 100, 5]
```


```python
from typing import Optional

def bubble_up(array: list[int|float], end: Optional[int]=None) -> None:
    # Get the length of the array
    n = len(array)
    # If end is not provided, process the entire array
    if end is None:
        end = n
    # Process the array up to the specified end index
    for i in range(end-1):
        # Compare and swap adjacent elements if they are in the wrong order
        if array[i] > array[i + 1]:
            array[i], array[i + 1] = array[i + 1], array[i]
                

```

<!-- grade_id: cs03, points: 1, grade: true, solution: false -->
```python
array: list[int|float] = [1, 100, 2, 3, 4, 5]
bubble_up(array)
assert array == [1, 2, 3, 4, 5, 100]

array: list[int|float] = [1, 100, 2, 3, 4, 5]
bubble_up(array, 5)
assert array == [1, 2, 3, 4, 100, 5]


```


**CS1-b.** Create a function that implements **Bubble Sort version 1** (from your Notes) to sort an array of integers. The function should sort the input array in place. Refer to the above pseudocode.

```python
def bubble_sort(array: list[int|float]) -> None:
    # Get the length of the array
    n = len(array)
    # Iterate over the array
    for big_idx in range(1, n):
        # Iterate over the array up to the current big_idx
        for small_idx in range(1, n):
            # Compare and swap adjacent elements if they are in the wrong order
            if array[small_idx - 1] > array[small_idx]:
                array[small_idx - 1], array[small_idx] = array[small_idx], array[small_idx - 1]

```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
array: list[int|float] = [10,3,8,47,1,0,-39,8,4,7,6,-5]
bubble_sort(array)
print(array)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]

```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS1-c.** Modify the code CS1-b so that it makes use of `bubble_up` function. Compare and discuss the difference in the performance.

```python
def another_bubble_sort(array: list[int|float]) -> None:
    # Get the length of the array
    n = len(array)
    # Iterate over the array
    for big_idx in range(1, n):
        # Use bubble_up function to move the largest element to the end
        bubble_up(array, n)
    return array

```

Output:
```
<cell>8: [1m[91merror:[0m No return value expected  [0m[93m[return-value][0m
```

<!-- grade_id: cs13, points: 1, grade: true, solution: false -->
```python
array: list[int|float] = [10,3,8,47,1,0,-39,8,4,7,6,-5]
another_bubble_sort(array)
print(array)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]


```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
```

**CS2.** Modify CS1-b, so that it returns  the number of comparisons that are performed.

Hint: To count the number of comparisons, you can create an integer variable which you *increment* just before the `if` statement. 

```python
def bubble_sort(array: list[int|float]) -> int:
    # Initialize the count of comparisons
    count: int = 0
    # Get the length of the array
    n = len(array)
    # Iterate over the array
    for big_idx in range(1, n):
        # Iterate over the array up to the current big_idx
        for small_idx in range(1, n):
            # Increment the comparison count
            count += 1
            # Compare and swap adjacent elements if they are in the wrong order
            if array[small_idx - 1] > array[small_idx]:
                array[small_idx - 1], array[small_idx] = array[small_idx], array[small_idx - 1]
    # Return the total number of comparisons
    return count
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
array: list[int|float] = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 121


```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 121
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3-b.** Create a function that implements **Insertion Sort** to sort an array of floats. The function should sort the input array in place. 

```python
def insertion_sort(array: list[int|float]) -> None:
    # Get the length of the array
    n = len(array)
    # Iterate over the array starting from the second element
    for big_idx in range(1, n):
        # Initialize the small_idx with the current big_idx
        small_idx = big_idx
        # Shift elements of the sorted segment to the right to find the correct position for array[big_idx]
        while small_idx > 0 and array[small_idx - 1] > array[small_idx]:
            # Swap the elements if they are in the wrong order
            array[small_idx - 1], array[small_idx] = array[small_idx], array[small_idx - 1]
            # Move to the previous element
            small_idx -= 1

```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
array = [10.3,3.8,8.4,47.1,1.0,0,-39.8,8.4,4.7,7.6,-6.5,-5.0]
insertion_sort(array)
print(array)
assert array == [-39.8, -6.5, -5.0, 0.0, 1.0, 3.8, 4.7, 7.6, 8.4, 8.4, 10.3, 47.1]


```

Output:
```
[-39.8, -6.5, -5.0, 0, 1.0, 3.8, 4.7, 7.6, 8.4, 8.4, 10.3, 47.1]
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###


```

**CS4.** Write a function `gen_random_int` that generates a list of integers from `0` to `n - 1` with its sequence randomly shuffled. The function should take in an integer `n`, denoting the number of integers to be generated. 

Hint: 
- You can use `random.shuffle(mylist)`  to shuffle the elements of a list called `mylist`.
- Refer to [Python's Random module](https://docs.python.org/3.7/library/random.html) for more details on how to use `shuffle()`.

```python
import random

def gen_random_int(number: int, seed: int) -> list[int]:
    # Initialize an empty list to store the integers
    result = []
    
    # Generate a list of integers from 0 to number-1
    for i in range(number):
        result += [i]
        
    # Shuffle the list using the provided seed for reproducibility
    random.Random(seed).shuffle(result)
    
    # Return the shuffled list
    return result

```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
output = gen_random_int(10, 100)
print(output)
assert output == [4, 0, 5, 9, 3, 1, 6, 8, 7, 2]


```

Output:
```
[4, 0, 5, 9, 3, 1, 6, 8, 7, 2]
```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###


```