# Week 1 Problem Set
## Homework



```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
Version 1.0.5
```

**HW0-b.** Write a function `palindrome(s)`to check if the string `s` is a [Palindrome](https://en.wikipedia.org/wiki/Palindrome). Use iteration for this solution.


```python
def palindrome(s: str) -> bool:
    # Iterate over the first half of the string
    for idx in range(len(s)//2):
        # Compare the character with its corresponding character from the end
        if s[idx] != s[len(s)-idx-1]:
            # If any character does not match, it's not a palindrome
            return False
    # If all characters match, it's a palindrome
    return True
```

<!-- grade_id: hw01, points: 1, grade: true, solution: false -->
```python
assert not palindrome("moon") 
assert palindrome("noon") 
assert palindrome("a a") 
assert palindrome("ada") 
assert not palindrome("ad a") 

```

<!-- grade_id: hw02, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW-1-a.** Reorder the correct sequence of code to implement bubble sort algorithm version 2. You can check the pseudocode in [Bubble Sort and Insertion Sort](https://data-driven-world.github.io/2023/notes/sorting-algorithms/bubble-sort). 

The function should ...

- **modify/mutate** an existing array into a sorted one
- **return** the number of comparisons made

Hint: To count the number of comparisons made you can simply increment an integer variable (counter) right before the `if` statement.

**HW1-b.** Create a function to implement **Bubble Sort Algorithm version 2**. 

```python
def bubble_sort(array: list[int|float]) -> int:
    count: int = 0  # Initialize the comparison counter
    n = len(array)  # Get the length of the array
    swapped = True  # Initialize the swapped flag to enter the while loop
    while swapped:  # Continue looping until no elements are swapped
        swapped = False  # Reset the swapped flag for the current pass
        for i in range(1, n):  # Iterate through the array
            count += 1  # Increment the comparison counter
            if array[i-1] > array[i]:  # Compare adjacent elements
                array[i-1], array[i] = array[i], array[i-1]  # Swap if they are in the wrong order
                swapped = True  # Set the swapped flag to True
    return count  # Return the total number of comparisons made
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
array: list[int|float] = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count: int = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 121

```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 121
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** The solution in HW1 can be improved!

The n-th pass places the n-th largest element into its final place (i.e. after 1 cycle, the (1st) largest element will be in its correct position).

Hence, we can reduce subsequent iterations by only considering the other (unsorted) `n-1` elements in the array. 

Implement the algorithm of **Bubble Sort version 4** under [Bubble Sort and Insertion Sort](https://data-driven-world.github.io/2023/notes/sorting-algorithms/bubble-sort). 


```python
def bubble_sort(array: list[int|float]) -> int:
    count: int = 0  # Initialize the comparison counter
    n = len(array)  # Get the length of the array
    swapped = True  # Initialize the swapped flag to enter the while loop
    while swapped:  # Continue looping until no elements are swapped
        swapped = False  # Reset the swapped flag for the current pass
        new_n = 0  # Initialize the new length of the array to be sorted
        for i in range(1, n):  # Iterate through the array
            first_number = array[i-1]  # Get the first number to compare
            second_number = array[i]  # Get the second number to compare
            count += 1  # Increment the comparison counter
            if first_number > second_number:  # Compare adjacent elements
                array[i-1], array[i] = array[i], array[i-1]  # Swap if they are in the wrong order
                swapped = True  # Set the swapped flag to True
                new_n = i  # Update the new length of the array to be sorted
        n = new_n  # Reduce the length of the array to be sorted
    return count  # Return the total number of comparisons made

```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
array: list[int|float] = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count: int = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 66

```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 66
```

<!-- grade_id: None, points: 0, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```