
# Week 3 Problem Set

## Homeworks

```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
Version 1.0.5
```


**HW1.** *Fibonaci:* Write a function to find the Fibonacci number given an index. These are the example of a the first few numbers in Fibonacci series.

| Indices    | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7  | 8  | 9  |
|------------|---|---|---|---|---|---|---|----|----|----|
| The series | 0 | 1 | 1 | 2 | 3 | 5 | 8 | 13 | 21 | 34 |

We can write that the Fibonacci number at index *i* is given by:

$fibo(i) = fibo(i-1) + fibo(i-2)$

Use recursion for your implementation.

Write and test the code below.

```python
def fibonacci(index: int) -> int:
    # Base case: if index is 0, return 0
    if index == 0:
        return 0
    # Base case: if index is 1, return 1
    elif index == 1:
        return 1
    # Recursive case: return the sum of the two preceding numbers
    else:
        return fibonacci(index-1) + fibonacci(index-2)
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(3) == 2
assert fibonacci(7) == 13
assert fibonacci(9) == 34

```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Max-Heapify:* Recall the algorithm for max-heapify in restoring the *max-heap property* of a binary heap tree. In the previous implementation, we used iteration. Implement the same algorithm using recursion. 

```python
def left_of(index: int) -> int:
    # Return the index of the left child
    return 2 * index + 1

def right_of(index: int) -> int:
    # Return the index of the right child
    return 2 * index + 2

def max_heapify(array: list, index: int, heap_size: int) -> None:
    # Initialize largest as root
    largest = index
    left = left_of(index)
    right = right_of(index)
    
    # Check if left child exists and is greater than root
    if left < heap_size and array[left] > array[largest]:
        largest = left
    
    # Check if right child exists and is greater than the largest so far
    if right < heap_size and array[right] > array[largest]:
        largest = right
    
    # Change root if needed
    if largest != index:
        array[largest], array[index] = array[index], array[largest]
        
        # Recursively heapify the affected sub-tree
        max_heapify(array, largest, heap_size)
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
result: list[int] = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
result: list[int] = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
max_heapify(result, 1, len(result))
print(result)
assert result == [4, 16, 10, 14, 7, 9, 3, 2, 8, 1]

```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
[4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *String Permutation:* Write two functions to return an array of all the permutations of a string. For example, if the input is `abc`, the output should be

```
output = ["abc", "acb", "bac", "bca", "cab", "cba"]
```

The first function only has one argument which is the input string, i.e. `permutate(s)`. The second function `permutate_recursive(s1, s2)` is the recursive function with two arguments String 1 (`s1`) and and String 2 (`s2`). The first function calls the second function `permutate_recursive("", s)` at the beginning. The second function should use a loop to move a character from `str2` to `str1` and recursively invokes it with a new `str1` and `str2`. 

```python
def permutate(s: str) -> list[str]:
    # Initialize the result list
    result = []
    # Call the recursive function with initial values
    result = permutate_recursive("", s)
    return result

def permutate_recursive(s1: str, s2: str) -> list[str]:
    # Base case: if s2 is empty, return s1 as a single-element list
    if len(s2) == 0:
        return [s1]
    else:
        results = []
        # Loop through each character in s2
        for i in range(len(s2)):
            # Move character from s2 to s1
            new_s1 = s1 + s2[i]
            new_s2 = s2[:i] + s2[i+1:]
            # Recursively call the function with new values
            results.extend(permutate_recursive(new_s1, new_s2))
        return results
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
assert set(permutate("abc"))==set(["abc", "acb", "bac", "bca", "cab", "cba"])
assert set(permutate_recursive("", "abc"))==set(["abc", "acb", "bac", "bca", "cab", "cba"])
assert set(permutate("abcd"))==set(["abcd", "abdc", "acbd", "acdb", "adbc", "adcb", "bacd", "badc", "bcad", "bcda", "bdac", "bdca", "cabd", "cadb", "cbad", "cbda", "cdab", "cdba", "dabc", "dacb", "dbac", "dbca", "dcab", "dcba"])
assert set(permutate_recursive("", "abcd"))==set(["abcd", "abdc", "acbd", "acdb", "adbc", "adcb", "bacd", "badc", "bcad", "bcda", "bdac", "bdca", "cabd", "cadb", "cbad", "cbda", "cdab", "cdba", "dabc", "dacb", "dbac", "dbca", "dcab", "dcba"])

```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *GCD:* Implement the recursive implementation of Euclid's algorithm for Greatest Common Divisor (GCD).

Hint: [Wikipedia: Greatest Common Divisor](https://en.wikipedia.org/wiki/Greatest_common_divisor#Euclid's_algorithm)

```python
def gcd(a: int, b: int) -> int:
    # Base case: if a is 0, return b
    if a == 0:
        return b
    # Base case: if b is 0, return a
    if b == 0:
        return a
    # Base case: if a and b are equal, return a (or b)
    if a == b:
        return a
    else:
        # Recursive case: subtract the smaller number from the larger number
        if a > b:
            return gcd(a - b, b)
        else:
            return gcd(a, b - a)
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
assert gcd(0,0) == 0
assert gcd(2,0) == 2
assert gcd(0,2) == 2
assert gcd(2,3) == 1
assert gcd(4, 28) == 4
assert gcd(1024, 526) == 2

```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW5.** Implement  a simple `quicksort(array)` function. In this simple quicksort function, a pivot is chosen from the array. In our case we will choose the median of the input array as the pivot. Any number smaller than the pivot will be collected into the "left" array while any number greater than the pivot will be collected into the "right" array. The left and the right array is then sorted again using quicksort function. The final sorted array is given as `left + pivot + right`. 

The function should return a new sorted array instead of modifying the array in place.

Reference: 
- [Quicksort](https://en.wikipedia.org/wiki/Quicksort)

```python
def quicksort(array: list) -> list:
    # Base case: if the array has 1 or 0 elements, it is already sorted
    if len(array) <= 1:
        return array
    
    # Choose the pivot as the median element of the array
    pivot = array[len(array) // 2]
    
    # Partition the array into three parts: left, middle, and right
    left = [x for x in array if x < pivot]
    middle = [x for x in array if x == pivot]
    right = [x for x in array if x > pivot]
    
    # Recursively sort the left and right parts and concatenate the results
    return quicksort(left) + middle + quicksort(right)
```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
array: list[int] = [3,6,8,10,1,2,1]
assert quicksort(array) == sorted(array)
array: list[int] = [3,3,3,6,8,10,1,2,1]
assert quicksort(array) == sorted(array)

```

<!-- grade_id: hw52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```