
# Week 3 Problem Set

## Cohort Sessions

```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
Version 1.0.5
```


**CS1.** You have implemented factorial using iteration in Problem Set 1. Now, implement the factorial problem using a recursion. The function should takes in an Integer input and returns and Integer output which is the factorial of the input. Recall that:

$n! = n \times (n-1) \times (n-2) \times \ldots \times 2 \times 1$

You should consider the case when $n$ is zero and one as well. 

Write and test the code below.

```python
# recursion
def factorial_recursion(n: int) -> int:
    # Base case: if n is 0 or 1, return 1
    if n == 1 or n == 0:
        return 1
    else:
        # Recursive case: n * factorial of (n-1)
        return n * factorial_recursion(n - 1)
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert factorial_recursion(0) == 1
assert factorial_recursion(1) == 1
assert factorial_recursion(5) == 120
assert factorial_recursion(7) == 5040
assert factorial_recursion(11) == 39916800

```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** *Helper Function:* Write a function `palindrome(s)`to check if the string `s` is a [Palindrome](https://en.wikipedia.org/wiki/Palindrome). To do this, write another function `is_palindrome(s, left, right)` where `left` and `right` are indices from the left and from the right to check the character in `str`. Use recursion instead of iteration in this problem.


```python
def palindrome(s: str) -> bool:
    # Get the length of the string
    n = len(s)
    # Call the helper function is_palindrome with initial indices
    return is_palindrome(s, 0, n - 1)

def is_palindrome(s: str, left: int, right: int) -> bool:
    # Base case: if left index is greater than or equal to right index, it's a palindrome
    if left >= right:
        return True
    # If characters at left and right indices are not the same, it's not a palindrome
    if s[left] != s[right]:
        return False
    # Recursive case: move towards the center of the string
    return is_palindrome(s, left + 1, right - 1)
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
assert not palindrome("moon") 
assert palindrome("noon") 
assert palindrome("a a") 
assert palindrome("ada") 
assert not palindrome("ad a") 

```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Towers of Hanoi:* Write a function `move_disks(n, from_tower, to_tower, aux_tower)` which returns an array of String for the movement of disks that solves the [Towers of Hanoi problem](https://en.wikipedia.org/wiki/Tower_of_Hanoi). 
* The first argument `n` is an Integer input that gives information on the number of disk. 
* The second argument `from_tower` is a String which is the label of the origin tower.
* The third argument `to_tower` is a String which is the label of the destination tower.
* The last argument `aux_tower` is a String which is the label of the auxilary tower.


```python
def move_disks(n: int, from_tower: str, to_tower: str, aux_tower: str) -> list[str]:
    # Base case: only one disk to move
    if n == 1:
        return ["Move disk 1 from " + from_tower + " to " + to_tower + "."]
    else:
        result: list[str] = []
        # Move n-1 disks from from_tower to aux_tower using to_tower as auxiliary
        result += move_disks(n-1, from_tower, aux_tower, to_tower)
        # Move the nth disk from from_tower to to_tower
        result.append("Move disk " + str(n) + " from " + from_tower + " to " + to_tower + ".")
        # Move the n-1 disks from aux_tower to to_tower using from_tower as auxiliary
        result += move_disks(n-1, aux_tower, to_tower, from_tower)
        return result
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
result: list[str] = move_disks(3, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to B.", "Move disk 2 from A to C.", "Move disk 1 from B to C.", "Move disk 3 from A to B.", "Move disk 1 from C to A.", "Move disk 2 from C to B.", "Move disk 1 from A to B."]

result: list[str] = move_disks(4, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B.", "Move disk 3 from A to C.", "Move disk 1 from B to A.", "Move disk 2 from B to C.", "Move disk 1 from A to C.", "Move disk 4 from A to B.", "Move disk 1 from C to B.", "Move disk 2 from C to A.", "Move disk 1 from B to A.", "Move disk 3 from C to B.", "Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B."]


```

Output:
```
['Move disk 1 from A to B.', 'Move disk 2 from A to C.', 'Move disk 1 from B to C.', 'Move disk 3 from A to B.', 'Move disk 1 from C to A.', 'Move disk 2 from C to B.', 'Move disk 1 from A to B.']
['Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.', 'Move disk 3 from A to C.', 'Move disk 1 from B to A.', 'Move disk 2 from B to C.', 'Move disk 1 from A to C.', 'Move disk 4 from A to B.', 'Move disk 1 from C to B.', 'Move disk 2 from C to A.', 'Move disk 1 from B to A.', 'Move disk 3 from C to B.', 'Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.']
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** *Merge Sort:* Write functions to implement the Merge Sort algorithm. The first function `mergesort(array)` should takes in an array of Integers in `array`. The function should sort the array in place. The second function `merge(array, p, q, r)` should implements the merge procedure. This function takes in an array of Integers in. `array`, the starting index for the left array `p`, the ending index for the left array `q`, and the ending index for the right array `r`. You can use a helper function for your recursion if needed.

```python
def merge(array: list, p: int, q: int, r: int) -> None:
    # Calculate the sizes of the two subarrays to be merged
    nleft = q - p + 1
    nright = r - q
    
    # Create temporary arrays for left and right subarrays
    left_array = array[p:q + 1]
    right_array = array[q + 1:r + 1]
    
    # Initialize pointers for left, right, and destination arrays
    left = 0
    right = 0
    dest = p
    
    # Merge the temporary arrays back into the original array
    while left < nleft and right < nright:
        if left_array[left] <= right_array[right]:
            array[dest] = left_array[left]
            left += 1
        else:
            array[dest] = right_array[right]
            right += 1
        dest += 1
    
    # Copy any remaining elements of left_array, if any
    while left < nleft:
        array[dest] = left_array[left]
        left += 1
        dest += 1
    
    # Copy any remaining elements of right_array, if any
    while right < nright:
        array[dest] = right_array[right]
        right += 1
        dest += 1

def mergesort_recursive(array: list, p: int, r: int) -> None:
    # Base case: if the array has one or no elements, it is already sorted
    if p < r:
        # Find the middle point to divide the array into two halves
        q = (p + r) // 2
        
        # Recursively sort the first and second halves
        mergesort_recursive(array, p, q)
        mergesort_recursive(array, q + 1, r)
        
        # Merge the sorted halves
        merge(array, p, q, r)

def mergesort(array: list) -> None:
    # Call the recursive merge sort function with initial indices
    mergesort_recursive(array, 0, len(array) - 1)

```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
input_array: list[int] = [5, 2, 4, 7, 1, 3, 2, 6]
mergesort(input_array)
assert input_array == [1, 2, 2, 3, 4, 5, 6, 7]

```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```