# Week 5 Problem Set

## Homework

```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
Version 1.0.5
```

```python
from typing import TypeAlias
from typing import Optional, Any, Iterator
from __future__ import annotations

Number: TypeAlias = int | float
NumberList: TypeAlias = list[int|float] 
```

**HW1.** Modify the class `TurtleWorld` to include the following attribute and methods:
- `movement_queue` which is an attribute of the type `Queue` to store the movement list. **Each entry of this queue is an object which has the name of the turtle (e.g. "t1") and the movement list (e.g. "ulrr"). The choice of this object is left to you, e.g. it could be a tuple (`("t1", "ulrr")`) or a dictionary with the same information or your own custom class.
- `add_movement(turtle, movement)` which adds turtle movement to the queue `movement_queue` to be run later. The argument `turtle` is a string containing the turtle's name. The argument `movement` is another string for the movement. For example, value for `turtle` can be something like `'t1'` while the value for the `movement` can be something like `'uullrrdd'`.
- `run()` which executes all the movements in the queue.

```python
import math

class Coordinate:
    
    def __init__(self, x:Number=0, y:Number=0) -> None:
        # Initialize the x and y coordinates
        self.x = x
        self.y = y
        
    @property
    def distance(self) -> float:
        # Calculate the distance from the origin (0, 0) to the coordinate
        return math.sqrt(self.x * self.x + self.y * self.y)
    
    def __str__(self) -> str:
        # Return a string representation of the coordinate
        return f"({self.x}, {self.y})"
```

```python
class Queue:
    def __init__(self) -> None:
        # Initialize an empty list to store queue items
        self.__items: list[Any] = []
    
    def enqueue(self, item: Any) -> None:
        # Add an item to the end of the queue
        self.__items.append(item)

    def dequeue(self) -> Any:
        # Remove and return the item from the front of the queue
        if self.is_empty:
            return None
        return self.__items.pop(0)
    
    def peek(self) -> Any:
        # Return the item at the front of the queue without removing it
        if self.is_empty:
            return None
        return self.__items[0]
    
    @property
    def is_empty(self) -> bool:
        # Check if the queue is empty
        return len(self.__items) == 0
    
    @property
    def size(self) -> int:
        # Return the number of items in the queue
        return len(self.__items)
```

```python
# Class definition
class RobotTurtle:
    # Attributes:
    def __init__(self, name: str, speed: int=1) -> None:
        # Initialize the name, speed, and position of the turtle
        self.name: str = name
        self.speed: int = speed
        self._pos: Coordinate = Coordinate(0, 0)
        
    @property
    def name(self) -> str:
        # Getter for the name attribute
        return self._name
    
    @name.setter
    def name(self, value: str) -> None:
        # Setter for the name attribute
        if isinstance(value, str) and value != "":
            self._name: str = value
            
    @property
    def speed(self) -> int:
        # Getter for the speed attribute
        return self._speed
    
    @speed.setter
    def speed(self, value: int) -> None:
        # Setter for the speed attribute
        if isinstance(value, int) and value > 0:
            self._speed: int = value

    @property
    def pos(self) -> Coordinate:
        # Getter for the position attribute
        return self._pos
    
    def move(self, direction: str) -> None:
        # Move the turtle in the specified direction
        update: dict[str, Coordinate] = {'up' : Coordinate(self.pos.x, self.pos.y + self.speed),
                                         'down' : Coordinate(self.pos.x, self.pos.y - self.speed),
                                         'left' : Coordinate(self.pos.x - self.speed, self.pos.y),
                                         'right' : Coordinate(self.pos.x + self.speed, self.pos.y)}
        self._pos = update[direction]

    def tell_name(self) -> None:
        # Print the name of the turtle
        print(f"My name is {self.name}")

```

```python
class TurtleWorld:
    # Set of valid movements
    valid_movements: set[str] = set('udlr')
    # Mapping of movement characters to directions
    movement_map: dict[str, str] = {'u': 'up', 'd': 'down', 'l': 'left', 'r': 'right'}
    
    def __init__(self) -> None:
        # Dictionary to store turtles by their names
        self.turtles: dict[str, RobotTurtle] = {}
        # Queue to store movements
        self.movement_queue = Queue()
        
    def add_movement(self, turtle: str, movement: str) -> None:
        # Add a movement to the queue if the turtle exists and the movement is valid
        if turtle in self.turtles:
            valid_movement = ''.join([m for m in movement if m in self.valid_movements])
            if valid_movement:
                self.movement_queue.enqueue((turtle, valid_movement))
    
    def run(self) -> None:
        # Execute all movements in the queue
        while not self.movement_queue.is_empty:
            turtle, movement = self.movement_queue.dequeue()
            self.move_turtle(turtle, movement)
        
    def move_turtle(self, name: str, movement: str) -> None:
        # Move the turtle according to the movement string
        if name in self.turtles:
            for move in movement:
                self.turtles[name].move(self.movement_map[move])
    
    def add_turtle(self, name: str, speed: int) -> None:
        # Add a new turtle to the world
        self.turtles.update({name: RobotTurtle(name, speed)})
        
    def remove_turtle(self, name: str) -> None:
        # Remove a turtle from the world
        self.turtles.pop(name)
        
    def list_turtles(self) -> list[str]:
        # List all turtles in the world
        return sorted(list(self.turtles.keys()))

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
world: TurtleWorld = TurtleWorld()
assert isinstance(world.movement_queue, Queue)

world.add_turtle('t1', 1)
world.add_turtle('t2', 2)
world.add_movement('t1', 'ur')
world.add_movement('t2', 'urz')
assert str(world.turtles['t1'].pos) == '(0, 0)'
assert str(world.turtles['t2'].pos) == '(0, 0)'
assert world.movement_queue.size == 2

world.run()
assert str(world.turtles['t1'].pos) == '(1, 1)'
assert str(world.turtles['t2'].pos) == '(2, 2)'

world.add_movement('t1', 'ur')
world.add_movement('t2', 'urz')

world.run()
assert str(world.turtles['t1'].pos) == '(2, 2)'
assert str(world.turtles['t2'].pos) == '(4, 4)'


###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** Implement a radix sorting machine. A radix sort for base 10 integers is a *mechanical* sorting technique that utilizes a collection of bins:
- one main bin 
- 10 digit-bins

Each bin acts like a *queue* and maintains its values in the order that they arrive. The algorithm works as follows:
- it begins by placing each number in the main bin. 
- Then it considers each value digit by digit. The first value is removed from the main bin and placed in a digit-bin corresponding to the digit being considered. For example, if the ones digit is being considered, 534 will be placed into digit-bin 4 and 667 will placed into digit-bin 7. 
- Once all the values are placed into their corresponding digit-bins, the values are collected from bin 0 to bin 9 and placed back in the main bin (in that order). 
- The process continues with the tens digit, the hundreds, and so on. 
- After the last digit is processed, the main bin will contain the values in ascending order.

Create a class `RadixSort` that takes in a List of Integers during object instantiation. The class should have the following properties:
- `items`: is a List of Integers containing the numbers.

It should also have the following methods:
- `sort()`: which returns the sorted numbers from `items` as an `list` of Integers.
- `max_digit()`: which returns the maximum number of digits of all the numbers in `items`. For example, if the numbers are 101, 3, 1041, this method returns 4 as the result since the maximum digit is four from 1041. 
- `convert_to_str(items)`: which returns items as a list of Strings (instead of Integers). This function should pad the higher digits with 0 when converting an Integer to a String. For example if the maximum digit is 4, the following items are converted as follows. From `[101, 3, 1041]` to `["0101", "0003", "1041"]`.

Hint: Your implementation should make use of the generic `Queue` class, which you created, for the bins.

```python
class RadixSort:
    
    def __init__(self, my_list: list[int]) -> None:
        # Initialize the list of items to be sorted
        self.items = my_list
    
    def max_digit(self) -> int:
        # Return the maximum number of digits in any item
        return max(len(str(item)) for item in self.items)
    
    def convert_to_str(self, items: list[int]) -> list[str]:
        # Convert items to strings, zero-padded to the length of the longest item
        max_len = self.max_digit()
        return [f"{item:0{max_len}d}" for item in items] # 0{max_len}d: zero-padded to max_len characters as a decimal integer
    
    def sort(self) -> list[int]:
        # Perform radix sort on the items
        max_len = self.max_digit()
        bins = [Queue() for i in range(10)]  # Create 10 bins for each digit
        
        # Process each digit from least significant to most significant
        for digit in range(max_len - 1, -1, -1):
            # Place items into bins based on the current digit
            for item in self.convert_to_str(self.items):
                bins[int(item[digit])].enqueue(int(item))
            
            # Collect items from bins and update the list
            self.items = []
            for bin in bins:
                while not bin.is_empty:
                    self.items.append(bin.dequeue())
        
        return self.items
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
list1: RadixSort = RadixSort([101, 3, 1041])
assert list1.items == [101,3,1041]
assert list1.max_digit() == 4
assert list1.convert_to_str(list1.items) == ["0101", "0003", "1041"]
ans: list[int] = list1.sort()
print(ans)
assert ans == [3, 101, 1041]
list2: RadixSort = RadixSort([23, 1038, 8, 423, 10, 39, 3901])
assert list2.sort() == [8, 10, 23, 39, 423, 1038, 3901]
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Output:
```
[3, 101, 1041]
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** Write a class called `EvaluateFraction` that evaluates postfix notation implemented using Stack and Queue data structures. Postfix notation is a way of writing expressions without using parenthesis. For example, the expression `(1+2)*3` would be written as `1 2 + 3 *`. The class `EvaluateFraction` has the following method:
- `input(inp)`: which pushes the input input one at a time. For example, to create a postfix notation `1 2 + 3 *`, we can call this method repetitively, e.g. `e.input('1'); e.input('2'); e.input('+'); e.input('3'); e.input('*')`. Notice that the input is of String data type. 
- `evaluate()`: which returns the output of the expression.
- `get_fraction(inp)`: which takes in an input string and returns a `Fraction` object. 

Postfix notation is evaluated using a Stack. The input streams from `input()` are stored in a Queue. If the output of the Queue is a number, the item is pushed into the stack. If it is an operator, we will apply the operator to the two top most item n the stacks and push the result back into the stack. 

```python
class Stack:
    def __init__(self) -> None:
        # Initialize an empty list to store stack items
        self.__items: list[Any] = []

    def push(self, item: Any) -> None:
        # Add an item to the top of the stack
        self.__items.append(item)

    def pop(self) -> Any:
        # Remove and return the item from the top of the stack
        if not self.is_empty:
            return self.__items.pop()
        return None

    def peek(self) -> Any:
        # Return the item at the top of the stack without removing it
        if not self.is_empty:
            return self.__items[-1]
        return None

    @property
    def is_empty(self) -> bool:
        # Check if the stack is empty
        return len(self.__items) == 0

    @property
    def size(self) -> int:
        # Return the number of items in the stack
        return len(self.__items)
```

```python
class Queue:
    def __init__(self) -> None:
        # Initialize two stacks to implement the queue
        self.left_stack: Stack = Stack()
        self.right_stack: Stack = Stack()
    
    def enqueue(self, item: Any) -> None:
        # Add an item to the end of the queue (push onto left stack)
        self.left_stack.push(item)

    def dequeue(self) -> Any:
        # Remove and return the item from the front of the queue
        if self.right_stack.is_empty:
            # Transfer all items from left stack to right stack if right stack is empty
            while not self.left_stack.is_empty:
                self.right_stack.push(self.left_stack.pop())
        return self.right_stack.pop()

    def peek(self) -> Any:
        # Return the item at the front of the queue without removing it
        if self.right_stack.is_empty:
            # Transfer all items from left stack to right stack if right stack is empty
            while not self.left_stack.is_empty:
                self.right_stack.push(self.left_stack.pop())
        return self.right_stack.peek()
    
    @property
    def is_empty(self) -> bool:
        # Check if the queue is empty
        return self.left_stack.is_empty and self.right_stack.is_empty
        
    @property
    def size(self) -> int:
        # Return the number of items in the queue
        return self.left_stack.size + self.right_stack.size

```

```python
def gcd(a: int, b: int) -> int:
    # Calculate the greatest common divisor using the Euclidean algorithm
    if b == 0:
        return a
    else:
        return gcd(b, a % b)


class Fraction:
    def __init__(self, num: int, den: int) -> None:
        # Initialize the numerator and denominator
        self.num = num
        self.den = den if den != 0 else 1
    
    @property
    def num(self) -> int:
        # Getter for the numerator
        return self._num
    
    @num.setter
    def num(self, val: int) -> None:
        # Setter for the numerator
        if isinstance(val, (int, float)):
            self._num = int(val)
        else:
            print("Numerator must be an integer or float")
    
    @property
    def den(self) -> int:
        # Getter for the denominator
        return self._den
    
    @den.setter
    def den(self, val: int) -> None:
        # Setter for the denominator
        if isinstance(val, (int, float)):
            val = int(val)
            if val == 0:
                self._den = 1
            else:
                self._den = val
        else:
            print("Denominator must be an integer or float")
    
    def __str__(self) -> str:
        # Return a string representation of the fraction
        return f"{self.num}/{self.den}"    
    
    def simplify(self) -> Fraction:
        # Simplify the fraction
        common_divisor = gcd(self.num, self.den)
        return Fraction(self.num // common_divisor, self.den // common_divisor)
    
    def __add__(self, other) -> Fraction:
        # Add two fractions
        new_num = self.num * other.den + other.num * self.den
        new_den = self.den * other.den
        return Fraction(new_num, new_den).simplify()
        
    def __eq__(self, other) -> bool:
        # Check if two fractions are equal
        simp_self = self.simplify()
        simp_other = other.simplify()
        return (simp_self.den == simp_other.den and simp_self.num == simp_other.num)
    
    def __sub__(self, other) -> Fraction:
        # Subtract two fractions
        new_num = self.num * other.den - other.num * self.den
        new_den = self.den * other.den
        return Fraction(new_num, new_den).simplify()
    
    def __mul__(self, other) -> Fraction:
        # Multiply two fractions
        new_num = self.num * other.num
        new_den = self.den * other.den
        return Fraction(new_num, new_den).simplify()
    
    def __lt__(self, other) -> bool:
        # Check if one fraction is less than another
        return self.num * other.den < other.num * self.den
    
    def __le__(self, other) -> bool:
        # Check if one fraction is less than or equal to another
        return self.num * other.den <= other.num * self.den
    
    def __gt__(self, other) -> bool:
        # Check if one fraction is greater than another
        return self.num * other.den > other.num * self.den
    
    def __ge__(self, other) -> bool:
        # Check if one fraction is greater than or equal to another
        return self.num * other.den >= other.num * self.den

```

```python
class EvaluateFraction:

    operands: str = "0123456789"
    operators: str = "+-*/"
    
    def __init__(self) -> None:
        # Initialize the expression queue and the stack
        self.expression: Queue = Queue()
        self.stack: Stack = Stack()
    
    def input(self, item: str) -> None:
        # Enqueue the input item to the expression queue
        self.expression.enqueue(item)
    
    def evaluate(self) -> Fraction:
        # Evaluate the expression in postfix notation
        while not self.expression.is_empty:
            item = self.expression.dequeue()
            if '/' in item:
                # If the item is a fraction, push it to the stack
                self.stack.push(self.get_fraction(item))
            elif item in self.operators:
                # If the item is an operator, pop two operands from the stack and apply the operator
                op2 = self.stack.pop()
                op1 = self.stack.pop()
                result = self.process_operator(op1, op2, item)
                self.stack.push(result)
        # Return the final result from the stack
        return self.stack.pop()
    
    def get_fraction(self, inp: str) -> Fraction:
        # Convert the input string to a Fraction object
        num, den = map(int, inp.split('/'))
        return Fraction(num, den)
    
    def process_operator(self, op1: Fraction, op2: Fraction, op: str) -> Fraction:
        # Apply the operator to the two operands and return the result
        if op == '+':
            return op1 + op2
        elif op == '-':
            return op1 - op2
        elif op == '*':
            return op1 * op2
        elif op == '/':
            return Fraction(op1.num * op2.den, op1.den * op2.num).simplify()
        else:
            raise ValueError(f"Unknown operator: {op}")

```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
pe: EvaluateFraction = EvaluateFraction()
pe.input("1/2")
pe.input("2/3")
pe.input("+")
assert pe.evaluate()==Fraction(7, 6)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1/6")
pe.input("-")
assert pe.evaluate()==Fraction(1, 1)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1/6")
pe.input("-")
pe.input("3/4")
pe.input("*")
assert pe.evaluate()==Fraction(3, 4)
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** Modify HW2 so that it can work with MixedFraction. Write a class called `EvaluateMixedFraction` as a subclass of `EvaluateFraction`. You need to override the following methods:
- `get_fraction(inp)`: This function should be able to handle string input for MixedFraction such as `1 1/2` or `3/2`. It should return a `MixedFraction` object.
- `evaluate()`: This function should return `MixedFraction` object rather than `Fraction` object. 

```python
class MixedFraction(Fraction):
    def __init__(self, top: int, bot: int, whole: int=0) -> None:
        # Initialize the mixed fraction by converting it to an improper fraction
        num = whole * bot + top
        super().__init__(num, bot)

    def get_three_numbers(self) -> tuple[int, int, int]:
        # Return the mixed fraction as a tuple of (top, bot, whole)
        whole = self.num // self.den
        top = self.num % self.den
        bot = self.den
        return (top, bot, whole)

    def __str__(self) -> str:
        # Return a string representation of the mixed fraction
        top, bot, whole = self.get_three_numbers()
        if whole == 0:
            return f"{top}/{bot}"
        else:
            return f"{whole} {top}/{bot}"

```

```python
class EvaluateMixedFraction(EvaluateFraction):
    def get_fraction(self, inp: str) -> Fraction:
        # Check if the input contains a whole number part
        if ' ' in inp:
            whole, frac = inp.split(' ')
            num, den = map(int, frac.split('/'))
            return MixedFraction(num, den, int(whole))
        else:
            num, den = map(int, inp.split('/'))
            return MixedFraction(num, den)
    
    def evaluate(self) -> MixedFraction:
        # Evaluate the expression in postfix notation
        while not self.expression.is_empty:
            item = self.expression.dequeue()
            if '/' in item:
                # If the item is a fraction, push it to the stack
                self.stack.push(self.get_fraction(item))
            elif item in self.operators:
                # If the item is an operator, pop two operands from the stack and apply the operator
                op2 = self.stack.pop()
                op1 = self.stack.pop()
                result = self.process_operator(op1, op2, item)
                self.stack.push(result)
        # Return the final result from the stack as a MixedFraction
        answer = self.stack.pop()
        return MixedFraction(answer.num, answer.den)
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
pe: EvaluateMixedFraction = EvaluateMixedFraction()
pe.input("3/2")
pe.input("1 2/3")
pe.input("+")
out: MixedFraction = pe.evaluate() 
assert out == MixedFraction(1, 6, 3)
assert isinstance(out, MixedFraction)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1 1/8")
pe.input("-")
assert pe.evaluate() == MixedFraction(1, 24)

pe.input("1 1/2")
pe.input("2 2/3")
pe.input("+")
pe.input("1 1/6")
pe.input("-")
pe.input("5/4")
pe.input("*")
assert pe.evaluate() == MixedFraction( 3, 4, 3)
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW5.** Write a function that takes in a graph representation in a dictionary and convert it to an object-oriented representation using the `Graph` class. Refer to your cohort problem sets for the `Graph` and `Vertex` class definition. 

The function takes in a dictionary that has the following format:
```python
graph = {'vertex1': [neighbour1, neighbour2, ...],
         'vertex2': [neighbour1, ...]}
```

Notes:
- The graph is represented as and adjecancy list. 
- The key of the dictionary is the `id` of the vertex in the graph. 
- The value of the dictionary is a list of neighbours of that vertex.
- The element in the list of neighbours is the `id` of the neighbouring vertices. 
- The data type for `id` is a string. 

Example:
```python
graph = {"A": ["B", "C"], 
         "B": ["C", "D"],
         "C": ["D"],
         "D": ["C"], 
         "E": ["F"],
         "F": ["C"]}
```

```python
class Vertex:
    def __init__(self, id_: str="") -> None:
        # Initialize the vertex with an id and an empty dictionary of neighbours
        self.id_: str = id_
        self.neighbours: dict[Vertex, Number] = {}
    
    def add_neighbour(self, nbr_vertex: Vertex, weight: Number=0) -> None:
        # Add a neighbour vertex with an optional weight
        self.neighbours[nbr_vertex] = weight
    
    def get_neighbours(self) -> list[Vertex]:
        # Return a list of neighbour vertices
        return list(self.neighbours.keys())
    
    def get_weight(self, neighbour: Vertex) -> Optional[Number]:
        # Return the weight of the edge to a neighbour vertex
        return self.neighbours.get(neighbour, None)
    
    def __eq__(self, other) -> bool:
        # Check if two vertices are equal based on their id
        return self.id_ == other.id_
    
    def __lt__(self, other) -> bool:
        # Compare two vertices based on their id
        return self.id_ < other.id_
    
    def __hash__(self) -> int:
        # Return the hash of the vertex based on its id
        return hash(self.id_)
    
    def __str__(self) -> str:
        # Return a string representation of the vertex and its neighbours
        neighbours_ids = ', '.join([nbr.id_ for nbr in self.neighbours])
        return f"Vertex {self.id_} is connected to: {neighbours_ids}"
```

```python
class Graph:
    def __init__(self) -> None:
        # Initialize the graph with an empty dictionary of vertices
        self.vertices: dict[str, Vertex] = {}
        
    @property
    def num_vertices(self) -> int:
        # Return the number of vertices in the graph
        return len(self.vertices)

    def _create_vertex(self, id_: str) -> Vertex:
        # Create a new vertex with the given id
        return Vertex(id_)

    def add_vertex(self, id_: str) -> None:
        # Add a vertex to the graph if it doesn't already exist
        if id_ not in self.vertices:
            self.vertices[id_] = self._create_vertex(id_)

    def get_vertex(self, id_: str) -> Optional[Vertex]:
        # Get the vertex with the given id
        return self.vertices.get(id_, None)

    def add_edge(self, start_v: str, end_v: str, weight: Number=0) -> None:
        # Add an edge from start_v to end_v with an optional weight
        if start_v not in self.vertices:
            self.add_vertex(start_v)
        if end_v not in self.vertices:
            self.add_vertex(end_v)
        self.vertices[start_v].add_neighbour(self.vertices[end_v], weight)

    def get_neighbours(self, id_: str) -> list[str]:
        # Get the list of neighbours for the vertex with the given id
        vertex = self.get_vertex(id_)
        if vertex:
            return [nbr.id_ for nbr in vertex.get_neighbours()]
        return []
    
    def __contains__(self, val: str) -> bool:
        # Check if a vertex with the given id exists in the graph
        return val in self.vertices.keys()
    
    def __iter__(self):
        # Iterate over the vertices in the graph
        for k,v in self.vertices.items():
            yield v 

```

```python
def create_graph_object(g: dict[str, list[str]]) -> Graph:
    # Initialize an empty Graph object
    output: Graph = Graph()
    
    # Iterate through each vertex and its neighbours in the input dictionary
    for vertex, neighbours in g.items():
        # Add the vertex to the graph
        output.add_vertex(vertex)
        # Add edges from the vertex to each of its neighbours
        for neighbour in neighbours:
            output.add_edge(vertex, neighbour)
    
    # Return the constructed graph
    return output
```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
graph: dict[str, list[str]] = {"A": ["B", "C"], 
         "B": ["C", "D"],
         "C": ["D"],
         "D": ["C"], 
         "E": ["F"],
         "F": ["C"]}

output: Graph = create_graph_object(graph)
assert output.num_vertices == 6
assert sorted([x.id_ for x in output]) == ["A", "B", "C", "D", "E", "F"]
assert sorted(output.get_neighbours("A")) == ["B", "C"]
assert sorted(output.get_neighbours("C")) == ["D"]
assert sorted(output.get_neighbours("E")) == ["F"]
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: None, points: 0, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```